Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Lab 1: Introduction to Python 

## Exercise 1: Defining Variables (8 pts)  
1. Create 4 scalar variables `a,b,c,d` containing the values 1,2,3,4.  

<!-- grade_id: cell-f025eb3cb228233c, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

2. Define a 3-component horizontal vector `H` containing `a,b,c`.  

<!-- grade_id: cell-2cefcd5c9026e8d7, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

3. Define a 3-component vertical vector `V` containing `b,c,d`. 

_Note_: "Vertical vector" means the array should be of shape (3,1). 

<!-- grade_id: cell-f0b9d3b5a0fa2b32, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

4. Define a 2x2 matric `M` containing a,b,c,d as:  
$$ M = \begin{pmatrix} a & b \\ c & d \end{pmatrix} $$  

<!-- grade_id: Q1-4, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

5. Displace the values of:  
  a. The third element of `H`  
  b. The first element of `V`  
  c. The first row of `M`  
  d. The second column of `M`

<!-- grade_id: cell-1009a72856227599, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Exercise  2: Functions (7 points) 

Write a function that calculates the distance between two vectors in the form $R(\vec x, \vec y)$. Where $R$ is a function that takes two vectors (i.e. $\vec x$,$\vec y$) as input a returns a scalar value of distance. Use your function to calculate the distance between the follow vector pairs:  
  1. `x = [3, 5]`, `y = [3, 4]`  
  2. `x = [2, 0]`, `y = [0, 2]`  
  3. `x = [-100, 30]`, `y = [-10, -50]` 
  
Hint: [mathematical formula to use](http://mathonline.wikidot.com/the-distance-between-two-vectors).

<!-- grade_id: Q2-function, points: 4, grade: true, solution: true -->
```python
def R(x,y):
    '''Returns the distance between two vectors of length 2. 

            Parameters:
                    x (array): A vector of length 2
                    y (array): Another vector of length 2

            Returns:
                    distance (float): The distance between the two input vectors
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: Q2-1, points: 1, grade: true, solution: false -->
```python
# Q2.1 Autograde test (DO NOT EDIT)
assert R(np.array([3,5]),np.array([3,4])) == 1.0
```

<!-- grade_id: Q2-2, points: 1, grade: true, solution: false -->
```python
# Q2.2 Autograde test (DO NOT EDIT)
assert np.round(R(np.array([2,0]),np.array([0,2])),4) == 2.8284
```

<!-- grade_id: Q2-3, points: 1, grade: true, solution: false -->
```python
# Q2.3 Autograde test (DO NOT EDIT)
assert np.round(R(np.array([-100,30]),np.array([-10,-50])),4) == 120.4159
```

#### Bonus : (3 points)  
Can you "vectorize" (i.e. use element-wise vector operations to simplify) this function? 

## Exercise 3: Loading Data and Plotting (10 pts)  

1. Load the `well_data.txt` data using `numpy`'s `loadtxt` function. The data is formatted such that the columns are:  

  Depth (ft), Electrical Conductivity (mS/m), Temperature (C)   
 

<!-- grade_id: Q3-1, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

 2. Convert depth to meters.

<!-- grade_id: Q3-2, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

3. Produce the following plot using this data (you have creative freedom to make your plot extra  
beautiful). Make sure your figure has:  
  a. Title          (2 points; one for each subplot)  
  b. Axis labels    (2 points; one for each subplot)  
  d. Readable text  (2 points; one for each subplot)  

<!-- grade_id: Q3-3, points: 6, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Exercise 4: Loading Data and Plotting (10 pts)  

Load the data file `ammonoids.mat`, which contains the variables `NumAm` and `MasAm`. The values within `NumAm` represent the number of ammonoid shells found in each grid cell, while `MasAm` gives the average weight of the shells found in each grid cell. `NumAm` and `MasAm` are both $m \times n$ matrices of the same dimension (representing the grid in which the shells were originally found). In this case, we will define that there are $m$ cells in the $x$ direction with a spacing ($\Delta x$) of 10m, and $n$ cells in the $y$ direction with a spacing ($\Delta y$) of 10m (see figure).


```python
import numpy as np 
import scipy.io as sio

data  = sio.loadmat('ammonoids.mat')
NumAm = data['NumAm'].astype(np.float32)
MasAm = data['MasAm'].astype(np.float32)
```

Provide a brief (one or two sentence) about how I am loading in the data. For example what type of data structure is the variable `data`? What is the point of the `.astype()` call?

<!-- grade_id: cell-22599f6ff8d88405, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

a. Write an `if` statement that checks if the sizes of `MasAm` and `NumAm` are equal. If they are, set $m$ to the number of cells in the $x$-direction and $n$ to the number of cells in the $y$-direction. Hint: use `.shape` attribute of NumPy arrays to determine the arrays size. (4 pts)

<!-- grade_id: Q4-a, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

b. Create a new matrix `TMA` whose elements represent the total mass of the ammonoids within each grid cell. (1 pt)  

<!-- grade_id: Q4-b, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

c. Find the total and average mass of ammonoids per cell in the first row using two different methods: 1) using indexing and the `.mean()` method of NumPy arrays and 2) with a `for` loop. (8 pts)

<!-- grade_id: Q4-c, points: 8, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

d. Make a new variable called `TMA300` that is the same size as `TMA` using the `np.zeros(m,n)` function. If the total mass of the ammonoids in a cell is greater than 300, place the total mass (value in `TMA`) in your new variable at the same cell location. 

_Hint_: Use Boolean Indexing here!

<!-- grade_id: Q4-d, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

e. Repeat step (d) making a new variable `TMA100` that only includes mass values less than 100, and new variable `TMA200` that only include mass values greater than 100 and less than 300. (5 pts)

_Hint_: You can combine boolean indexing condition by enclosing the conditions in parentheses `(`_condition 1_`)`  and join them via an or(`|`) statement or and (`&`). i.e. `(`_condition 1_`)`  `|` `(`_condition 2_`)`

<!-- grade_id: Q4-e, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

f. Calculate the mean mass of the ammonoids greater than 100 and less than 300. Hint: use the function `np.nanmean`. (3 pts)

<!-- grade_id: Q4-f, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

g. Make a spatial plot of `TMA100`, `TMA200`, and `TMA300`. Hint: Make the spatial plot using `plt.imshow()`. Put all three plots on figure using the array of subplot axes `ax`, which can be indexed. Furthermore, use the `ax[i].set_title()` method to add a title to each subplot. 

<!-- grade_id: Q4-g, points: 9, grade: true, solution: true -->
```python
fig, ax = plt.subplots(1,3,sharex=True,sharey=True,figsize=(9,3))

# YOUR CODE HERE
raise NotImplementedError()

# Don't worry about this. Formatting a colorbar with subplots can be tricky.
# So I've done it for you, but can you explain what's going on?
plt.tight_layout()
plt.subplots_adjust(wspace=0.075,hspace=0.1,right=0.9)

cmap  = plt.get_cmap('viridis')            
cNorm = plt.matplotlib.colors.Normalize(vmin=0, vmax=400.0)
axcb  = fig.add_axes([0.925, 0.05, 0.03, 0.9])
cb1   = plt.matplotlib.colorbar.ColorbarBase(axcb, cmap=cmap, norm=cNorm)
```